// Numbas version: finer_feedback_settings
{"name": "Compose two permutations given in cycle notation", "extensions": ["permutations"], "custom_part_types": [{"source": {"pk": 30, "author": {"name": "Christian Lawson-Perfect", "pk": 1}, "edit_page": "/part_type/30/edit"}, "name": "Permutation in cycle notation", "short_name": "write-a-permutation-in-cycle-notation", "description": "<p>The student must write a permutation in cycle notation. You can optionally require the cycles to be disjoint and/or all of length 2.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "if(settings[\"must_be_transpositions\"],\n  as_transpositions(settings[\"correct_answer\"]),\n  show(settings[\"correct_answer\"])\n)", "hint": {"static": false, "value": "if(settings[\"must_be_disjoint\"],\n  if(settings[\"must_be_transpositions\"],\n    \"Write a permutation in cycle notation as a product of disjoint transpositions.\",\n    \"Write a permutation in cycle notation as a product of disjoint cycles.\"\n  ),\n  if(settings[\"must_be_transpositions\"],\n    \"Write a permutation in cycle notation as a product of transpositions.\",\n    \"Write a permutation in cycle notation.\"\n  )\n)"}, "allowEmpty": {"static": true, "value": false}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\napply(student_permutation);\napply(is_correct);\napply(is_disjoint);\napply(is_transpositions)\n\ninterpreted_answer:\nstudent_permutation\n\nis_correct:\ncorrectif(student_permutation=settings[\"correct_answer\"])\n\nis_disjoint:\nif(is_disjoint(studentAnswer),\n  true\n,\n  if(settings[\"must_be_disjoint\"],\n    set_credit(0,\"You have not written your answer as a product of disjoint cycles.\")\n  );\n  false\n)\n\nstudent_permutation:\ntry(\n  permutation(studentAnswer),\n  message,\n    warn(\"Your answer is not valid cycle notation: \"+message);\n    fail(\"Your answer is not valid cycle notation: \"+message)\n)\n\nis_transpositions:\nif(is_transpositions(studentAnswer),\n  true\n,\n  if(settings[\"must_be_transpositions\"],\n    set_credit(0,\"You have not written your answer as a product of transpositions.\")\n  );\n  false\n)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "apply(student_permutation);\napply(is_correct);\napply(is_disjoint);\napply(is_transpositions)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "student_permutation"}, {"name": "is_correct", "description": "", "definition": "correctif(student_permutation=settings[\"correct_answer\"])"}, {"name": "is_disjoint", "description": "", "definition": "if(is_disjoint(studentAnswer),\n  true\n,\n  if(settings[\"must_be_disjoint\"],\n    set_credit(0,\"You have not written your answer as a product of disjoint cycles.\")\n  );\n  false\n)"}, {"name": "student_permutation", "description": "", "definition": "try(\n  permutation(studentAnswer),\n  message,\n    warn(\"Your answer is not valid cycle notation: \"+message);\n    fail(\"Your answer is not valid cycle notation: \"+message)\n)"}, {"name": "is_transpositions", "description": "", "definition": "if(is_transpositions(studentAnswer),\n  true\n,\n  if(settings[\"must_be_transpositions\"],\n    set_credit(0,\"You have not written your answer as a product of transpositions.\")\n  );\n  false\n)"}], "settings": [{"name": "correct_answer", "label": "Correct answer", "help_url": "", "hint": "", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "must_be_disjoint", "label": "Must the cycles be disjoint?", "help_url": "", "hint": "", "input_type": "checkbox", "default_value": false}, {"name": "must_be_transpositions", "label": "Must the cycles be transpositions?", "help_url": "", "hint": "", "input_type": "checkbox", "default_value": false}], "public_availability": "always", "published": true, "extensions": ["permutations"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Compose two permutations given in cycle notation", "tags": [], "metadata": {"description": "<p>The student is given two random non-identity in $S_n$ (n=10 by default), in cycle notation. They must compose them and write the result in cycle notation.</p>", "licence": null}, "statement": "<p>\\begin{align}<br/>\\pi_1 &amp;= \\var{perm1} \\\\<br/>\\pi_2 &amp;= \\var{perm2}<br/>\\end{align}</p>\n<p></p>", "advice": "", "rulesets": {}, "extensions": ["permutations"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"n": {"name": "n", "group": "Ungrouped variables", "definition": "10", "description": "<p>The size of the $S_n$ this question works in.</p>", "templateType": "anything", "can_override": true}, "perm1": {"name": "perm1", "group": "Ungrouped variables", "definition": "permutation(deal(n))", "description": "", "templateType": "anything", "can_override": false}, "perm2": {"name": "perm2", "group": "Ungrouped variables", "definition": "permutation(deal(n))", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "perm1<>permutation([]) and perm2<>permutation([])", "maxRuns": 100}, "ungrouped_variables": ["n", "perm1", "perm2"], "variable_groups": [], "functions": {"cycle": {"parameters": [["cycle", "list of number"], ["n", "number"]], "type": "anything", "language": "javascript", "definition": "var o = [];\nfor(var i=0;i<n;i++) {\n  var j = cycle.indexOf(i);\n  if(j==-1) { \n    o.push(i); \n  } else {\n    o.push(cycle[(j+1)%cycle.length]);\n  }\n}\nreturn o;"}}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "write-a-permutation-in-cycle-notation", "useCustomName": false, "customName": "", "marks": "2", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Write $\\pi_1\\pi_2$ in cycle notation.</p>", "alternatives": [{"type": "write-a-permutation-in-cycle-notation", "useCustomName": true, "customName": "$\\pi_2\\pi_1$", "marks": 0, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "alternativeFeedbackMessage": "<p>It looks like you calculated $\\pi_2\\pi_1$.</p>", "useAlternativeFeedback": false, "settings": {"correct_answer": "perm2*perm1", "show_hint": true, "must_be_disjoint": false, "must_be_transpositions": false, "identity_name": "e"}}], "settings": {"correct_answer": "perm1*perm2", "show_hint": true, "must_be_disjoint": false, "must_be_transpositions": false, "identity_name": "e"}}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "http://clppc.local:8000/accounts/profile/1/"}, {"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Stefan Kolb", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/590/"}, {"name": "Laura Midgley", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/18287/"}, {"name": "Shweta Sharma", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/21418/"}], "resources": []}]}], "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "http://clppc.local:8000/accounts/profile/1/"}, {"name": "Christian Lawson-Perfect", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/7/"}, {"name": "Stefan Kolb", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/590/"}, {"name": "Laura Midgley", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/18287/"}, {"name": "Shweta Sharma", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/21418/"}]}