// Numbas version: finer_feedback_settings
{"name": "Graph theory custom part type", "extensions": ["graph-theory"], "custom_part_types": [{"source": {"pk": 22, "author": {"name": "Christian Lawson-Perfect", "pk": 1}, "edit_page": "/part_type/22/edit"}, "name": "Graph editor", "short_name": "graph-editor", "description": "<p>Manipulate or edit a graph.</p>", "help_url": "", "input_widget": "graph-editor", "input_options": {"correctAnswer": "settings[\"initial_graph\"]", "hint": {"static": true, "value": ""}, "graph": {"static": false, "value": "settings[\"initial_graph\"]"}, "mode": {"static": false, "value": "settings[\"mode\"]"}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\nfeedback(draw(studentAnswer));\ncorrect()\n\ninterpreted_answer:\nstudentAnswer", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "feedback(draw(studentAnswer));\ncorrect()"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "studentAnswer"}], "settings": [{"name": "initial_graph", "label": "Initial graph", "help_url": "", "hint": "", "input_type": "code", "default_value": "graph([])", "evaluate": true}, {"name": "mode", "label": "Mode", "help_url": "", "hint": "", "input_type": "dropdown", "default_value": "edit movethings", "choices": [{"value": "edit movethings", "label": "Edit"}, {"value": "play movethings", "label": "Move things"}, {"value": "play hamilton", "label": "Make a Hamiltonian cycle"}, {"value": "play euler", "label": "Make an Eulerian cycle"}]}], "public_availability": "restricted", "published": false, "extensions": ["graph-theory"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Graph theory custom part type", "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "", "advice": "", "rulesets": {}, "extensions": ["graph-theory"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true, "j": false}, "constants": [], "variables": {"g": {"name": "g", "group": "Ungrouped variables", "definition": "random_planar_graph(4)", "description": "", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["g"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "graph-editor", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "num_vertices: len(vertices(studentAnswer))\n\nmark:\n  assert(is_connected(studentAnswer),\n    fail(\"Your graph is not a single connected component.\")\n  );\n  if(is_tree(studentAnswer),\n     add_credit(0.5, \"Your graph is a tree.\")\n  ,  negative_feedback(\"Your graph is not a tree.\")\n  );\n  if(num_vertices=5,\n     add_credit(0.5, \"Your graph has 5 vertices.\")\n  ,  negative_feedback(\"Your graph has {num_vertices} vertices.\")\n  )", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Draw a tree with 5 vertices.</p>", "settings": {"initial_graph": "graph([])", "mode": "edit movethings"}}, {"type": "graph-editor", "useCustomName": false, "customName": "", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "mark:\n  correctif(adjacency_matrix(studentAnswer) = adjacency_matrix(g))", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Draw a graph with this adjacency matrix</p>\n<p>$\\var{adjacency_matrix(g)}$</p>", "settings": {"initial_graph": "graph([])", "mode": "edit movethings"}}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "http://clppc.local:8000/accounts/profile/1/"}], "resources": []}]}], "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "http://clppc.local:8000/accounts/profile/1/"}]}