// Numbas version: finer_feedback_settings
{"name": "Permutations - explore mode", "extensions": ["permutations"], "custom_part_types": [{"source": {"pk": 30, "author": {"name": "Christian Lawson-Perfect", "pk": 1}, "edit_page": "/part_type/30/edit"}, "name": "Permutation in cycle notation", "short_name": "write-a-permutation-in-cycle-notation", "description": "<p>The student must write a permutation in cycle notation. You can optionally require the cycles to be disjoint and/or all of length 2.</p>", "help_url": "", "input_widget": "string", "input_options": {"correctAnswer": "if(settings[\"must_be_transpositions\"],\n  as_transpositions(settings[\"correct_answer\"]),\n  show(settings[\"correct_answer\"])\n)", "hint": {"static": false, "value": "if(settings[\"must_be_disjoint\"],\n  if(settings[\"must_be_transpositions\"],\n    \"Write a permutation in cycle notation as a product of disjoint transpositions.\",\n    \"Write a permutation in cycle notation as a product of disjoint cycles.\"\n  ),\n  if(settings[\"must_be_transpositions\"],\n    \"Write a permutation in cycle notation as a product of transpositions.\",\n    \"Write a permutation in cycle notation.\"\n  )\n)"}, "allowEmpty": {"static": true, "value": false}}, "can_be_gap": true, "can_be_step": true, "marking_script": "mark:\napply(student_permutation);\napply(is_correct);\napply(is_disjoint);\napply(is_transpositions)\n\ninterpreted_answer:\nstudent_permutation\n\nis_correct:\ncorrectif(student_permutation=settings[\"correct_answer\"])\n\nis_disjoint:\nif(is_disjoint(studentAnswer),\n  true\n,\n  if(settings[\"must_be_disjoint\"],\n    set_credit(0,\"You have not written your answer as a product of disjoint cycles.\")\n  );\n  false\n)\n\nstudent_permutation:\ntry(\n  permutation(studentAnswer),\n  message,\n    warn(\"Your answer is not valid cycle notation: \"+message);\n    fail(\"Your answer is not valid cycle notation: \"+message)\n)\n\nis_transpositions:\nif(is_transpositions(studentAnswer),\n  true\n,\n  if(settings[\"must_be_transpositions\"],\n    set_credit(0,\"You have not written your answer as a product of transpositions.\")\n  );\n  false\n)", "marking_notes": [{"name": "mark", "description": "This is the main marking note. It should award credit and provide feedback based on the student's answer.", "definition": "apply(student_permutation);\napply(is_correct);\napply(is_disjoint);\napply(is_transpositions)"}, {"name": "interpreted_answer", "description": "A value representing the student's answer to this part.", "definition": "student_permutation"}, {"name": "is_correct", "description": "", "definition": "correctif(student_permutation=settings[\"correct_answer\"])"}, {"name": "is_disjoint", "description": "", "definition": "if(is_disjoint(studentAnswer),\n  true\n,\n  if(settings[\"must_be_disjoint\"],\n    set_credit(0,\"You have not written your answer as a product of disjoint cycles.\")\n  );\n  false\n)"}, {"name": "student_permutation", "description": "", "definition": "try(\n  permutation(studentAnswer),\n  message,\n    warn(\"Your answer is not valid cycle notation: \"+message);\n    fail(\"Your answer is not valid cycle notation: \"+message)\n)"}, {"name": "is_transpositions", "description": "", "definition": "if(is_transpositions(studentAnswer),\n  true\n,\n  if(settings[\"must_be_transpositions\"],\n    set_credit(0,\"You have not written your answer as a product of transpositions.\")\n  );\n  false\n)"}], "settings": [{"name": "correct_answer", "label": "Correct answer", "help_url": "", "hint": "", "input_type": "code", "default_value": "", "evaluate": true}, {"name": "must_be_disjoint", "label": "Must the cycles be disjoint?", "help_url": "", "hint": "", "input_type": "checkbox", "default_value": false}, {"name": "must_be_transpositions", "label": "Must the cycles be transpositions?", "help_url": "", "hint": "", "input_type": "checkbox", "default_value": false}], "public_availability": "always", "published": true, "extensions": ["permutations"]}], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false, "typeendtoleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Permutations - explore mode", "tags": [], "metadata": {"description": "", "licence": "None specified"}, "statement": "<p>This question is about the permutation</p>\n<p>\\[p = \\var{twoline(p)} \\]</p>", "advice": "", "rulesets": {}, "extensions": ["permutations"], "builtin_constants": {"e": false, "pi,\u03c0": true, "i": true, "j": false}, "constants": [{"name": "e", "value": "perm(\"\")", "tex": "e"}], "variables": {"n": {"name": "n", "group": "Setup", "definition": "12", "description": "", "templateType": "number", "can_override": false}, "p": {"name": "p", "group": "Setup", "definition": "permutation(deal(n))", "description": "", "templateType": "anything", "can_override": false}, "num_cycles": {"name": "num_cycles", "group": "Setup", "definition": "len(cycles(p))", "description": "", "templateType": "anything", "can_override": false}, "remainder": {"name": "remainder", "group": "Finding cycles", "definition": "p", "description": "", "templateType": "anything", "can_override": false}, "n_to_try": {"name": "n_to_try", "group": "Finding cycles", "definition": "if(remainder<>e, remainder_cycles[0][0]+1, 0)\n", "description": "<p>The first element in the next cycle to find.</p>", "templateType": "anything", "can_override": false}, "first_cycle": {"name": "first_cycle", "group": "Finding cycles", "definition": "if(remainder<>e,\n  cycle(cycles(remainder)[0]),\n  perm(\"\")\n)", "description": "", "templateType": "anything", "can_override": false}, "partial": {"name": "partial", "group": "Finding cycles", "definition": "p*inverse(remainder)", "description": "", "templateType": "anything", "can_override": false}, "remainder_cycles": {"name": "remainder_cycles", "group": "Finding cycles", "definition": "cycles(remainder)", "description": "", "templateType": "anything", "can_override": false}, "partial_reminder": {"name": "partial_reminder", "group": "Finding cycles", "definition": "\"<p>You have shown that $p = \\\\var{partial} \\\\cdot q$, for some $q$.</p>\\n<p>Let\\'s find another cycle in $q$.</p>\"", "description": "", "templateType": "long string", "can_override": false}, "cycle_reminder": {"name": "cycle_reminder", "group": "Finding cycles", "definition": "\"<p>You have shown that $p = \\\\var{partial}$.</p>\"", "description": "", "templateType": "long string", "can_override": false}, "first_cycle_prompt": {"name": "first_cycle_prompt", "group": "Finding cycles", "definition": "\"<p>Let\\'s find a cycle in $p$.</p>\"", "description": "", "templateType": "long string", "can_override": false}}, "variablesTest": {"condition": "num_cycles > 1", "maxRuns": 100}, "ungrouped_variables": [], "variable_groups": [{"name": "Setup", "variables": ["n", "p", "num_cycles"]}, {"name": "Finding cycles", "variables": ["remainder", "remainder_cycles", "n_to_try", "first_cycle", "partial", "partial_reminder", "cycle_reminder", "first_cycle_prompt"]}], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "numberentry", "useCustomName": true, "customName": "How many cycles?", "marks": "10", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Write down the cycles", "rawLabel": "Write down the cycles", "otherPart": 1, "variableReplacements": [], "availabilityCondition": "remainder<>e and credit<1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": "Calculate the number of cycles", "prompt": "<p>{if(partial&lt;&gt;perm(\"\"), cycle_reminder, \"\")}</p>\n<p>How many disjoint, non-trivial cycles are there in $p$?</p>", "minValue": "num_cycles", "maxValue": "num_cycles", "correctAnswerFraction": false, "allowFractions": false, "mustBeReduced": false, "mustBeReducedPC": 0, "displayAnswer": "", "showFractionHint": true, "notationStyles": ["plain", "en", "si-en"], "correctAnswerStyle": "plain"}, {"type": "write-a-permutation-in-cycle-notation", "useCustomName": true, "customName": "Find a cycle", "marks": "0", "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [{"label": "Find a cycle", "rawLabel": "", "otherPart": 1, "variableReplacements": [{"variable": "remainder", "definition": "remainder*inverse(interpreted_answer)"}], "availabilityCondition": "answered and credit=1 and len(cycles(remainder))>1", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": false}, {"label": "How many cycles?", "rawLabel": "", "otherPart": 0, "variableReplacements": [{"variable": "remainder", "definition": "remainder*inverse(interpreted_answer)"}], "availabilityCondition": "answered and credit=1 and len(cycles(remainder))<2", "penalty": "", "penaltyAmount": 0, "showPenaltyHint": true, "lockAfterLeaving": false}], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>{if(partial&lt;&gt;perm(\"\"), partial_reminder, first_cycle_prompt)}</p>\n<p>Write down the cycle containing the element $\\var{n_to_try}$.</p>", "settings": {"correct_answer": "first_cycle", "must_be_disjoint": true, "must_be_transpositions": false}}], "partsMode": "explore", "maxMarks": 0, "objectives": [{"name": "Calculate the number of cycles", "limit": "10"}], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "http://clppc.local:8000/accounts/profile/1/"}], "resources": []}]}], "contributors": [{"name": "Christian Lawson-Perfect", "profile_url": "http://clppc.local:8000/accounts/profile/1/"}]}